# Привет!

Перед началом прохождения убедитесь, что все настроено как надо. Следующая ячейка должна выполниться без ошибок:

```python
print("Are you ready?")

import numpy as np

print(np.__version__)
```

Output:
```
Are you ready?
ModuleNotFoundError                       Traceback (most recent call last)
Cell In [1], line 3
      1 print("Are you ready?")
----> 3 import numpy as np
      5 print(np.__version__)

ModuleNotFoundError: No module named 'numpy'
```

После прохождения не забудьте **перезапустить ядро** (**restart the kernel**) (в панели меню выберите Kernel $\rightarrow$ Restart) и затем **запуск всех ячеек** (**run all cells**) (в панели меню выберите Cell $\rightarrow$ Run All).

Убедитесь, что заменили все такие места 
```python
# YOUR CODE HERE
raise NotImplementedError()
```
и
```
"YOUR ANSWER HERE"
```
на свои решения и ответы.

Также, заполните свои данные (имя и группу):

```python
NAME = "Cюй Mинчао"
GROUP = "749512"
```

```python
import uuid

STUDENT_UUID = uuid.uuid4()
print(f"Your uuid is {STUDENT_UUID}")
```

---

# Приветствуем тебя, путник!

Это первая контрольная работа с некоторыми вопросами по теме numpy, pandas и python.

```python
import numpy as np
import pandas as pd

np.random.seed(123)
```

### Задание 1 (1 балл)

Напишите код для разворачивания строки и приведения её к верхнему регистру. 

Например, на вход подаётся строка "Привет", на выходе - "ТЕВИРП".

<!-- grade_id: cell-69a75a6a9f247b59, points: 0, grade: false, solution: true -->
```python
def reverse_string_tall_letters(string: str) -> str:
    s2 = string.upper()
    return s2[::-1]
    # raise NotImplementedError()
```

<!-- grade_id: cell-6c2c7d17e93c9be4, points: 1, grade: true, solution: false -->
```python
assert len(reverse_string_tall_letters("hi")) == 2 
assert reverse_string_tall_letters("x") == "X"
assert reverse_string_tall_letters("") == ""
assert reverse_string_tall_letters("Привет!") == "!ТЕВИРП"

```

Output:
```
NameError                                 Traceback (most recent call last)
C:\Users\MIKUSA~1\AppData\Local\Temp/ipykernel_3708/1232275509.py in <module>
----> 1 assert len(reverse_string_tall_letters("hi")) == 2
      2 assert reverse_string_tall_letters("x") == "X"
      3 assert reverse_string_tall_letters("") == ""
      4 assert reverse_string_tall_letters("Привет!") == "!ТЕВИРП"
      5 

NameError: name 'reverse_string_tall_letters' is not defined
```

### Задание 2 (1 балл)

Выведите уникальные значения из входного массива (матрицы).

<!-- grade_id: cell-a83f3275964bb308, points: 0, grade: false, solution: true -->
```python
def get_unique_rows(matrix: np.ndarray) -> np.ndarray:
    # YOUR CODE HERE
    uniques=np.unique(matrix,axis=0)
    return uniques
    raise NotImplementedError()
```

<!-- grade_id: cell-4d50a5c607725df3, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose(get_unique_rows(np.array([[1]])), [[1]])
np.testing.assert_allclose(get_unique_rows(np.array([[1], [1], [1]])), [[1]])
np.testing.assert_allclose(get_unique_rows(np.array([[1, 2], [1, 2], [2, 1]])), [[1, 2], [2, 1]])

```

### Задание 3 (2 балла)

Выведите N наименьших значений из массива (матрицы). 

> ⚠️ Если N больше, чем общее число элементов в массиве, то нужно вернуть весь исходный массив. 

<!-- grade_id: cell-d4aba3080a3cd06f, points: 0, grade: false, solution: true -->
```python
def get_n_smallest(arr: np.ndarray, n: int) -> np.ndarray:
    # YOUR CODE HERE
    a = list(arr.reshape(-1))
    a.sort()
    if len(a)<=n:
        return a
    ans = []
    for i in range(n):
        ans.append(a[i])
    return ans
    raise NotImplementedError()
```

<!-- grade_id: cell-922af51c5bbde513, points: 2, grade: true, solution: false -->
```python
assert len(get_n_smallest(np.array([2, 1, 3]), 1)) == 1
np.testing.assert_allclose(get_n_smallest(np.zeros(5), 6), [0]*5)

```

### Задание 4 (2 балла)

Найдите самое часто повторяющееся число в массиве. 

> ⚠️ Обратите внимание на аннотацию функции! Важно привести (cast) к нужному типу результат

<!-- grade_id: cell-db76e26eedba62d1, points: 0, grade: false, solution: true -->
```python
def get_freq_value(arr: np.ndarray) -> float:
    # YOUR CODE HERE
    a = list(arr)
    dir = {}
    for aa in a:
        if aa in dir.keys():
            dir[aa]+=1
        else:
            dir[aa] = 1
    max_key = 0
    max_val = 0
    for i in dir.keys():
        if dir[i]>max_val:
            max_key = i
            max_val = dir[i]
    return float(max_key)
    raise NotImplementedError()
```

<!-- grade_id: cell-6617d7118be7f720, points: 2, grade: true, solution: false -->
```python
assert isinstance(get_freq_value(np.random.randint(0, 10, 10)), float)
np.testing.assert_almost_equal(get_freq_value(np.array([1, 1])), 1.0)
np.testing.assert_almost_equal(get_freq_value(np.array([1, 3, 2, 2])), 2.0)

```

### Задание 5 (1 балл)

Преобразуйте два списка в pandas.Series, где первый списoк - это индексы, а второй - значения. 

<!-- grade_id: cell-fa84a84ced474d2f, points: 0, grade: false, solution: true -->
```python
def get_series(list_1: list, list_2: list) -> pd.Series:
    # YOUR CODE HERE
    s=pd.Series(list_2,list_1)
    return s
    raise NotImplementedError()
```

<!-- grade_id: cell-110c1c0a2650e2fa, points: 1, grade: true, solution: false -->
```python
assert isinstance(get_series(["a"], [1]), pd.Series)
np.testing.assert_allclose(get_series(["a"], [1]).shape, 1)
pd.testing.assert_index_equal(get_series(["a", "c"], [1, 3]).index, pd.Index(["a", "c"]))
np.testing.assert_allclose(get_series(["a", "c", "1"], [1, 4, 5]).values, [1, 4, 5])

```

### Задание 6 (3 балла)

Разделите одну колонку со строками на 2 колонки с названиями "STD" и "City". 

> ⚠️ Обратите внимание, что первая строка не является данными, а названия городов не должно содержать пробелов!

<!-- grade_id: cell-58dc34e723ed3302, points: 0, grade: false, solution: true -->
```python
def separate_columns(df: pd.DataFrame) -> pd.DataFrame: 
    # YOUR CODE HERE
    df_list = list(df['row'])
    for i in range(len(df_list)):
        if i == 0:continue
        df_list[i] = df_list[i].split(',')
    return pd.DataFrame(df_list[1:],columns=['STD','City'])
    raise NotImplementedError()
```

<!-- grade_id: cell-ddbd80fb8adda273, points: 3, grade: true, solution: false -->
```python
df = pd.DataFrame(
    [
        "STD, City",
        "33, Kolkata",
        "44,Chennai",
        "40, Hyderabad",
        "80,Bangalore",
        ", TricksTown"
    ],
    columns=['row']
)
df_result = separate_columns(df)

assert isinstance(df_result, pd.DataFrame)
assert set(df_result.columns) == set(["STD", "City"])
np.testing.assert_allclose(df_result.shape, (5, 2))
pd.testing.assert_series_equal(df_result.iloc[1], pd.Series(["44", "Chennai"]), check_index=False, check_names=False)

```

### Задание 7 (2 балла)

Выберите только те строки, в которых животное - это кошка, а возраст не меньше 2-х лет, не включая строки с NaN.

<!-- grade_id: cell-92c17c0ec05456d2, points: 0, grade: false, solution: true -->
```python
def animals_selection(df: pd.DataFrame) -> pd.DataFrame:
    # YOUR CODE HERE
    ans = []
    df_1 = df["animal"].tolist()
    df_2 = df['age'].tolist()
    df_l = []
    for j in range(len(df_1)):
        df_l.append([df_1[j],df_2[j]])
    for i in df_l:
        if i[0] == None or i[1] == None:continue
        if i[0] == 'cat' and int(i[1]>2):
            ans.append(i)
    df1 = pd.DataFrame(ans,columns=["animal", "age"])
    return df1
    raise NotImplementedError()
```

<!-- grade_id: cell-2c583ee44d239300, points: 2, grade: true, solution: false -->
```python
df = pd.DataFrame(
    zip(['cat', np.nan, 'cat', 'dog', 'cat', 'dog'], [1.95, 2, 3, 0.5, np.nan, 5]),
    columns=["animal", "age"]
)
df_result = animals_selection(df)

assert isinstance(df_result, pd.DataFrame)
np.testing.assert_allclose(df_result.shape, (1, 2))
np.testing.assert_allclose(df_result.age, [3])
assert not (df_result.animal.isnull()).all()

```

### Задание 8 (2 балла)

Напишите класс StatComputer, который имеет три метода:
- метод mean_cols(), который возвращает среднее значение каждой колонки в виде массива значений
- метод std_rows(), который возвращает стандартное отклонение каждой строки в виде массива значений
- метод sum_full(), который возвращает сумму всех элементов массива в виде числа

Конструктор класса должен принимать массив чисел и возвращать результаты вычислений по каждому методу.

<!-- grade_id: cell-955168e77c9346fa, points: 0, grade: false, solution: true -->
```python
class StatComputer:
    # YOUR CODE HERE
    def __init__(self,data):
        self.data = data
    def mean_cols(self):
        dataT = list(np.mean(self.data,axis=0))
        return dataT
    def std_rows(self):
        ans = np.std(self.data,axis=1)
        ans = list(ans.T)
        return ans
    def sum_full(self):
        return sum(sum(self.data))
    # raise NotImplementedError()
```

<!-- grade_id: cell-d52d4a43af9134f9, points: 2, grade: true, solution: false -->
```python
data = np.array([[3, 2, 3, 5, 5],
                 [0, 8, 5, 7, 0],
                 [6, 4, 3, 6, 3],
                 [1, 7, 9, 8, 9]])
stat_comp = StatComputer(data)

np.testing.assert_almost_equal(stat_comp.mean_cols(), [2.5 , 5.25, 5., 6.5, 4.25])
np.testing.assert_almost_equal(stat_comp.std_rows(), [1.2, 3.4058773, 1.356466, 2.9933259])
np.testing.assert_almost_equal(stat_comp.sum_full(), 94)

```

### Задание 9 (3 балла)

Напишите функцию, которая вернет количество столбцов с медианным значением больше задаваемого порога `threshold` после матричного умножения матриц `mul1` и `mul2`.

> ⚠️ Главное, вспомните, что такое "матричное умножение"

<!-- grade_id: cell-4941db9c0f636e2b, points: 0, grade: false, solution: true -->
```python
def high_median_n_cols_after_matrix_multiplication(mul1: np.array, mul2: np.array, threshold: float) -> int:
    # YOUR CODE HERE
    c = np.matmul(mul1, mul2)
    zhon = list(np.median(c,axis=0))
    ans = 0
    for z in zhon:
        if z>threshold:
            ans+=1
    return ans
    # raise NotImplementedError()
```

<!-- grade_id: cell-1b72bf4632ef61bc, points: 3, grade: true, solution: false -->
```python
# Test 1
x1 = np.array([[5, 6, 4],
               [0, 5, 0]])
x2 = np.array([[6, 2],
               [8, 0],
               [1, 7]])
result = high_median_n_cols_after_matrix_multiplication(x1, x2, threshold=30)

np.testing.assert_equal(result, 1)

# Test 2
x1 = np.array([[11,  3],
               [ 7, 13],
               [ 8,  8],
               [ 3,  2],
               [11,  5]])
x2 = np.array([[12,  5,  3,  3],
               [13,  1,  0,  9]])
result = high_median_n_cols_after_matrix_multiplication(x1, x2, threshold=40)

np.testing.assert_equal(result, 3)

```

### Задание 10 (1 балл)

Опишите `broadcasting` механизм во фреймворке `numpy`. Приведите примеры.

> Для примеров кода можно использовать конструкцию внутри markdown ячеек:

```python
def code_sample():
    pass
```

Или создавать ячейки с кодом и писать в них.

<!-- grade_id: cell-37ca0151f6ea53e2, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

```python
import numpy as np
A = np.array([[56.0, 0.0, 4.4, 68.0],
[1.2, 104.0, 52.0, 8.0],
[1.8, 135.0, 99.0, 0.9]])

cal = np.sum(A,axis=0)
percent = A/cal
print(percent)
```